// intro1.rs
//
// 关于`I AM NOT DONE`这件事：
// 我们有时会鼓励您在给定的练习中继续尝试，即使您已经弄清楚了。 如果您一切正常并准备好
// 进行下一个练习，请删除下面的`I AM NOT DONE`注释。
//
// 如果您使用`rusdlings watch`运行此程序：当您更改以下行任一行时，练习文件将重新加载！
// 尝试添加 `println!` 行，或者尝试更改它在终端中的输出内容。 尝试删除一个分号，看看会发生什么！
//
// 执行 `rusdlingshint intro1` 或使用 `hint` watch 子命令获取提示。

// I AM NOT DONE

fn main() {
    println!("你好，还有");
    println!(r#"       欢迎来到…                           "#);
    println!(r#"                 _   _ _                  "#);
    println!(r#"  _ __ _   _ ___| |_| (_)_ __   __ _ ___  "#);
    println!(r#" | '__| | | / __| __| | | '_ \ / _` / __| "#);
    println!(r#" | |  | |_| \__ \ |_| | | | | | (_| \__ \ "#);
    println!(r#" |_|   \__,_|___/\__|_|_|_| |_|\__, |___/ "#);
    println!(r#"                               |___/      "#);
    println!();
    println!("本练习编译成功。 其余练习包含编译器或逻辑错误。 Rustlings 背后的核心概念是修复这些错误");
    println!("并解决练习。 祝你好运！");
    println!();
    println!("本练习的源代码位于“exercises/intro/intro1.rs”。 看一看！");
    println!(
        "展望未来，练习的源码将始终在成功/失败输出中。"
    );
    println!();
    println!(
        "如果您想使用 rust-analyzer（Rust 的 LSP 实现），请确保您的编辑器已设置，然后在继续之前"
    );
    println!("运行“rustlings lsp”。")
}
