// intro1.rs
//
// About this `I AM NOT DONE` thing:
// We sometimes encourage you to keep trying things on a given exercise, even
// after you already figured it out. If you got everything working and feel
// ready for the next exercise, remove the `I AM NOT DONE` comment below.
//
// If you're running this using `rustlings watch`: The exercise file will be
// reloaded when you change one of the lines below! Try adding a `println!`
// line, or try changing what it outputs in your terminal. Try removing a
// semicolon and see what happens!
//
// Execute `rustlings hint intro1` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    println!("你好，还有");
    println!(r#"       欢迎来到…                           "#);
    println!(r#"                 _   _ _                  "#);
    println!(r#"  _ __ _   _ ___| |_| (_)_ __   __ _ ___  "#);
    println!(r#" | '__| | | / __| __| | | '_ \ / _` / __| "#);
    println!(r#" | |  | |_| \__ \ |_| | | | | | (_| \__ \ "#);
    println!(r#" |_|   \__,_|___/\__|_|_|_| |_|\__, |___/ "#);
    println!(r#"                               |___/      "#);
    println!();
    println!("本练习编译成功。 其余练习包含编译器或逻辑错误。 Rustlings 背后的核心概念是修复这些错误");
    println!("并解决练习。 祝你好运！");
    println!();
    println!("本练习的源代码位于“exercises/intro/intro1.rs”。 看一看！");
    println!(
        "展望未来，练习的源码将始终在成功/失败输出中。"
    );
    println!();
    println!(
        "如果您想使用 rust-analyzer（Rust 的 LSP 实现），请确保您的编辑器已设置，然后在继续之前"
    );
    println!("运行“rustlings lsp”。")
}
